# CS 445 Setup Check

Run every cell in this notebook (**Run → Run All Cells**). If it finishes
with no errors and prints `Setup looks good.`, your environment is working
and you are ready for the first assignment.


```python
import torch
import torchvision
import sklearn
import matplotlib
import pandas as pd
import nbgrader

```

<!-- grade_id: checkVersions, points: 1, grade: true, solution: false -->
```python
from packaging.version import Version

assert Version(torch.__version__) >= Version('2.12'), torch.__version__
assert Version(nbgrader.__version__) >= Version('0.9'), nbgrader.__version__

```

```python
# A plot should appear below this cell.
import matplotlib.pyplot as plt

plt.plot([0, 1, 2, 3], [0, 1, 4, 9], marker='o')
plt.title('If you can see this, plotting works')
plt.show()

print('Setup looks good.')

```